```python
# Initialize Otter
import otter
grader = otter.Notebook("PCE-02-Lists.ipynb")
```

# Programming Component Excercise 02 &ndash; Lists

## Q1: Create a list

### Q1.1 Autograded task

Write a function called `create_list` that takes three strings and an integer, puts them in a list in the order that they're received, and returns the list.

Parameters:
- `string1` : `str`

    Any string
    
    
- `string2` : `str`

    Any string


- `string3` : `str`

    Any string


- `num1` : `int`

    Any integer

Returns:

- `result` : `list` 

    A list that contains the three input strings, and an integer in order.


⬇️⬇️⬇️ Write your solution in the cell below. ⬇️⬇️⬇️

```python

def create_list(string1, string2, string3, num1):
    
    result = ...
    return result

# Try different arguments! Test your function in the cell below using these 3 strings. Feel free to change them to run your own experiments.
s1 = "apples"
s2 = "oranges"
s3 = "bananas"
n1 = 9001
create_list(s1, s2, s3, n1)
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q1")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

### Q1.2: Manually graded task

Go back to the solution cell above, and add a docstring to the `create_list` function. The docstring should contain:

1. A brief description of what the function does;
2. A list of the parameters with a brief description of each parameter and its type;
3. A brief description of the return value.

For this task you are free to format the docstring however you like. If you are unsure on how to get started, refer to the &ldquo;Functions&rdquo; slides on ELMS for an example of a docstring that includes all the required pieces of information.

__NOTE__: _There is no autograder for this task: we will grade it manually. Note that you can earn points for this task even if your solution for the previous task is not passing the autograder, however we will check that your docstring does not generate syntax errors of its own._

⬇️⬇️⬇️ You can use the cell below to see what your docstring looks like. ⬇️⬇️⬇️

```python
help(create_list)
```

## Q2: Find the length of a list

Write a function called `list_len` that tells you if a given list is a small list if length is less than or equal to 3 and big list otherwise.

Parameters:
- `input_list` : `list` 

    A list


Returns:

- `result` : `str` 

    A string that can take the following values:
    
    - `"small list"` if `input_list` has length 3 or less
    - `"big list"` otherwise

⬇️⬇️⬇️ Write your solution in the cell below. ⬇️⬇️⬇️

```python
def list_len(input_list):
    ...

# Test with different arguments!
day_list = ["Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday", "Sunday"]
list_len(day_list)
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q2")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q3: Retrieve the fourth item from list
Write a function `get_fourth_item` that returns the fourth item in a list.

Parameters:
- `items` : `list`

   A list of items, any type

Returns:
- `result` : (any type)

   The fourth element in the list

*hint: indexes (which you use to grab items from a list) and numerical positions (e.g., the fourth item) are not the same!*

⬇️⬇️⬇️ Write your solution in the cell below. ⬇️⬇️⬇️

```python
def get_fourth_item(items):
    ...
    
# Test with different arguments!
get_fourth_item(["INST126", "INST201", "INST311", "INST326", "INST335"])
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q3")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q4: Find smallest number

Write a function `find_smallest` that returns the smallest item from a list of numbers. For example, given the list `[77, 68, 100, 95]` as an argument, your function should return the value `68`.

Parameters:
- `nums` : `list`
        
    A list of numbers (each an `int`)

Returns:
- `result` : `int` 

    The smallest number in the list

To do this, you should:
1. Sort the list (either in descending or ascending order),
2. Retrieve the first (or the last, depending on order) item in the sorted list;


*Hint: remember that the `.sort()` method sorts the list in place and thus is non-fruitful (i.e., it returns `None`). This is unlike the `sorted()` function, which is fruitful and does not sort the list in place: it returns a new, sorted list instead.*

---

This problem is worth 2 points:
- 1 point for writing a function that returns correctly the smallest item.
- 1 point for using the sort function and index access instead of the `min()` function

---

⬇️⬇️⬇️ Write your solution in the cell below. ⬇️⬇️⬇️

```python
def find_smallest(nums):
    ...

# Test with different arguments!
find_smallest([77, 68, 100, 95, 86, 82, 99])
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q4")
```

```python

```

```python

```

### **_NOTE:_**  Do not forget to run this cell

```python
!jupyter nbconvert --to script PCE-02-Lists.ipynb
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['PCE-02-Lists.py'])
```

Tests for q1 (points: None)
```
>>> assert type(create_list('apples', 'oranges', 'bananas', 111)) == list, 'The return value should be a list'
```
```
>>> assert create_list('apples', 'oranges', 'bananas', 111) == ['apples', 'oranges', 'bananas', 111], 'The return value should be a list containing the values of s1, s2, s3, and n1'
```
```
>>> assert create_list('apples', 'oranges', 'kiwis', 118) == ['apples', 'oranges', 'kiwis', 118], 'The return value should be a list containing the values of s1, s2, s3, and n1'
```
```
>>> assert create_list('apples', 'lemons', 'durian', 113) == ['apples', 'lemons', 'durian', 113], 'The return value should be a list containing the values of s1, s2, s3, and n1'
```


Tests for q2 (points: None)
```
>>> assert type(list_len(['Mon', 'Tue', 'Wed', 'Thu', 'Fri', 'Sat', 'Sun'])) == str, 'The function should return a str'
```
```
>>> assert list_len(['Mon', 'Tue', 'Wed', 'Thu', 'Fri', 'Sat', 'Sun']) == 'big list', 'The function should return "big list" for the input ["Mon", "Tue", "Wed", "Thu", "Fri", "Sat", "Sun"]'
```
```
>>> assert list_len([0, 1, 2, 3]) == 'big list', 'The function should return "big lidt" for the input [0, 1, 2, 3]'
```
```
>>> assert list_len(['Friday', 'Saturday', 'Sunday']) == 'small list', 'The function should return "small list" for the input ["Friday", "Saturday", "Sunday"]'
```


Tests for q3 (points: None)
```
>>> assert get_fourth_item(['INST126', 'INST201', 'INST311', 'INST326', 'INST335']) == 'INST326', 'For input ["INST126", "INST201", "INST311", "INST326", "INST335"], the function should return "INST326"'
```
```
>>> assert get_fourth_item(['INST326', 'INST201', 'INST311', 'INST126', 'INST335']) == 'INST126', 'For input ["INST326", "INST201", "INST311", "INST126", "INST335"], the function should return "INST126"'
```
```
>>> assert get_fourth_item(['INST126', 'INST326', 'INST311', 'INST201', 'INST335']) == 'INST201', 'For input ["INST126", "INST326", "INST311", "INST201", "INST335"], the function should return "INST201"'
```
```
>>> assert get_fourth_item(['INST126', 'INST201', 'INST326', 'INST311', 'INST335']) == 'INST311', 'For input ["INST126", "INST201", "INST326", "INST311", "INST335"], the function should return "INST311"'
```


Tests for q4 (points: None)
```
>>> assert type(find_smallest([77, 68, 100, 95, 86, 82, 99])) == int, 'The function should return an int'
```
```
>>> assert find_smallest([77, 68, 100, 95, 86, 82, 99]) == 68, 'The function should return 68 for the input [77, 68, 100, 95, 86, 82, 99]'
```
```
>>> assert find_smallest([32, 76, 97, 54, 86]) == 32, 'The function should return 32'
```
```
>>> assert find_smallest([43, 88, 67, 77, 54, 91, 22]) == 22, 'The function should return 22'
```
