```python
# Initialize Otter
import otter
grader = otter.Notebook("PCE-01-Conditionals.ipynb")
```

# Programming Component Exercises 01 &ndash; Conditionals

## Q1: Test if number is even

Write a function called `is_even` that tests whether `num` is even: the function should `True` if the number is even, and `False` if it is not.

_Hint_: consider using the modulo operator (`%`) to determine is a number is divisible by two.

```python
...

# Try with different arguments!
is_even(num=10)
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q1")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q2: Speeding ticket algorithm

Write a function called `speeding_ticket` that takes an observed speed and a speed limit  (both in miles per hour) as parameters, and returns either the string `"ok"` (if the speed is at or below the limit), `"grace"` (if the speed is over, but no more than 10 mph over the limit), or `"ticket"` (if the speed is more than 10 mph over the limit).

The two parameters should be called `speed` and `limit`.

*Hint: you may find a chained conditional pattern helpful here.*

```python
def speeding_ticket(speed, limit):
    ...

# Test with different arguments!
speeding_ticket(speed=25, limit=35)
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q2")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Manually graded: Add a docstring to `speeding_ticket`

Go back to the solution cell above, and add a docstring to the `speeding_ticket` function. The docstring should contain:

1. A brief description of what the function does;
2. A list of the parameters with a brief description of each parameter and its type;
3. A brief description of the return value.

For this task you are free to format the docstring however you like. If you are unsure on how to get started, refer to the &ldquo;Functions&rdquo; slides on ELMS for an example of a docstring that includes all the required pieces of information.

__NOTE__: _There is no autograder for this task: we will grade it manually. Note that you can earn points for this task even if your solution for the previous task is not passing the autograder, however we will check that your docstring does not generate syntax errors of its own._

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q3: A complex Boolean expression

Write a function `complex_boolean_numfilter` that returns `True` when `x` is an even number falling OUTSIDE of the interval $[10, 50]$, `False` otherwise.

```python
def complex_boolean_numfilter(x):
    ...

# Test with different arguments!
complex_boolean_numfilter(x=52)
```

The cell below is an autograder cell, run it to check if your solution is correct.

```python
grader.check("q3")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

### **_NOTE:_**  Do not forget to run this cell

```python
!jupyter nbconvert --to script PCE-01-Conditionals.ipynb
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['PCE-01-Conditionals.py'])
```

Tests for q1 (points: None)
```
>>> assert type(is_even(num=2)) == bool, 'The function should return a boolean'
```
```
>>> assert is_even(num=2) == True, 'Wrong return value'
```
```
>>> assert is_even(num=1) == False, 'Wrong return value'
```
```
>>> assert is_even(num=0) == True, 'Wrong return value'
```
```
>>> assert is_even(num=-1) == False, 'Wrong return value'
```


Tests for q2 (points: None)
```
>>> assert type(speeding_ticket(speed=35, limit=55)) == str, 'The function should return a str'
```
```
>>> assert speeding_ticket(speed=50, limit=55) == 'ok', 'Wrong return value'
```
```
>>> assert speeding_ticket(speed=65, limit=55) == 'grace', 'Wrong return value'
```
```
>>> assert speeding_ticket(speed=73, limit=55) == 'ticket', 'Wrong return value'
```


Tests for q3 (points: None)
```
>>> assert type(complex_boolean_numfilter(x=2)) == bool, 'The function should return a bool'
```
```
>>> assert complex_boolean_numfilter(x=17) == False, 'Wrong return value'
```
```
>>> assert complex_boolean_numfilter(x=57) == False, 'Wrong return value'
```
```
>>> assert complex_boolean_numfilter(x=78) == True, 'Wrong return value'
```
```
>>> assert complex_boolean_numfilter(x=5) == False, 'Wrong return value'
```
```
>>> assert complex_boolean_numfilter(x=20) == False, 'Wrong return value'
```
