# ASSIGNMENT CONFIG
export_cell:
    pdf: false
    run_tests: true
generate: 
    filtering: true
    zips: false
    show_stdout: true

# PCE 01 &ndash; Expressions and Variables [Total: 5 pts]

# BEGIN QUESTION
name: q1

## Q1: Add two numbers together [1 pt]

Inside function `sumof2` write an expression that sums `x` to `y` and assigns the result in `z`. 

# BEGIN SOLUTION

```python
# BEGIN SOLUTION NO PROMPT
def sumof2(x, y):
    z = x + y
    return z
# END SOLUTION
""" # BEGIN PROMPT
def sumof2(x, y):
    z = ...
    return z
""" # END PROMPT

# Try different arguments!
sumof2(3, 4)
```

# END SOLUTION

The cell below is an autograder cell, run it to check if your solution is correct.

# BEGIN TESTS

```python
assert type(sumof2(3, 4)) == int, "The type of z should be int"
```

<!-- grade_id: cell-ee3bbef34aba60d3, points: 1, grade: true, solution: false -->
```python
assert sumof2(3, 4) == 7, "The value of z should be the sum of x and y"
```

```python
assert sumof2(10, 4) == 14, "The value of z should be the sum of x and y"
```

```python
assert sumof2(10, -1) == 9, "The value of z should be the sum of x and y"
```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# BEGIN QUESTION
name: q2

## Q2: Add the dollar sign at the end of a string [1 pt]

Inside function `add_dollar`, write an expression that adds the dollar sign at the end of variable `s`. 

# BEGIN SOLUTION

<!-- grade_id: cell-95f65ad4ca16ef7d, points: 0, grade: false, solution: true -->
```python
# BEGIN SOLUTION NO PROMPT
def add_dollar(s):
    s = s + '$'
    return s
# END SOLUTION
""" # BEGIN PROMPT
def adddollar(s):
    s = ...
    return s
""" # END PROMPT

# Try different arguments!
add_dollar("hello")
```

# END SOLUTION

# BEGIN TESTS

```python
assert type(add_dollar("hello")) == str, "The type of s should be string"
```

```python
assert add_dollar("hello") == "hello$", "The return value should be the concatenation of s and '$'"
```

```python
assert add_dollar("good") == "good$", "The return value should be the concatenation of s and '$'"
```

```python
assert add_dollar("song") == "song$", "The return value should be the concatenation of s and '$'"
```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# BEGIN QUESTION
name: q3

## Q3: Create a phrase with three words [1 pt]

Inside functions `phraseof3`, write an expression that joins `word1`, `word2`, and `word3`, including spaces in between, and assigns the result to `phrase`.

# BEGIN SOLUTION

<!-- grade_id: cell-38c8b454a131779e, points: 1, grade: true, solution: true -->
```python
# BEGIN SOLUTION NO PROMPT
def phraseof3(word1, word2, word3):
    phrase = word1 + " " + word2 + " " + word3
    return phrase
# END SOLUTION
""" # BEGIN PROMPT
def phraseof3(word1, word2, word3):
    phrase = ...
    return phrase
""" # END PROMPT

# Try different arguments!
phraseof3("Boys", "watch", "TV")
```

# END SOLUTION

# BEGIN TESTS

```python
assert type(phraseof3("Boys", "watch", "TV")) == str, "The type of phrase should be str"
```

```python
assert phraseof3("Boys", "watch", "TV") == "Boys watch TV", "Wrong value; make sure to check your expression"
```

```python
assert phraseof3("Cats", "love", "purring") == "Cats love purring", "Wrong value; make sure to check your expression"
```

```python
assert phraseof3("Jane", "eats", "pizza") == "Jane eats pizza", "Wrong value; make sure to check your expression"
```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# BEGIN QUESTION
name: q4

## Q4: Add (in the math sense) two numbers that are represented initially as *strings* [1 pt]

Inside function `add_as_ints` write an expression that, after converting them to integers, sums `x` to `y` and assigns the result to `z`.

# BEGIN SOLUTION

```python
# BEGIN SOLUTION NO PROMPT
def add_as_ints(x, y):
    z = int(x) + int(y)
    return z
# END SOLUTION
""" # BEGIN PROMPT
def add_as_ints(x, y):
    z = ...
    return z
""" # END PROMPT

# Try different arguments!
add_as_ints("55", "33")
```

# END SOLUTION

# BEGIN TESTS

```python
assert type(add_as_ints("55", "33")) == int, "The type of z should be int"
```

```python
assert add_as_ints("55", "33") == 88, "The value of z should be the sum of x and y (as ints)"
```

```python
assert add_as_ints("1", "2") == 3, "The value of z should be the sum of x and y (as ints)"
```

```python
assert add_as_ints("10", "-1") == 9, "The value of z should be the sum of x and y (as ints)"
```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# BEGIN QUESTION
name: q5

## Q5: Add (in the string sense) two strings that are represented initially as *numbers* [1 pt]

Inside function `add_as_strs` write an expression that, after converting them to strings, concatenates `x` with `y` and assigns the result to `z`.

# BEGIN SOLUTION

<!-- grade_id: cell-600a1fa05e2bbd12, points: 1, grade: true, solution: true -->
```python
# BEGIN SOLUTION NO PROMPT
def add_as_strs(x, y):
    z = str(x) + str(y)
    return z
# END SOLUTION
""" # BEGIN PROMPT
def add_as_strs(x, y):
    z = ...
    return z
""" # END PROMPT

# Try different arguments!
add_as_strs(55, 33)
```

# END SOLUTION

# BEGIN TESTS

```python
assert type(add_as_strs(55, 33)) == str, "The type of z should be str"
```

```python
assert add_as_strs(55, 33) == '5533', "The value of z should be the sum of x and y (as strings)"
```

```python
assert add_as_strs(1, 2) == '12', "The value of z should be the sum of x and y (as strings)"
```

```python
assert add_as_strs(10, -1) == '10-1', "The value of z should be the sum of x and y (as strings)"
```

# END TESTS

# END QUESTION