```python
# Initialize Otter
import otter
grader = otter.Notebook("PCE-02-Strings.ipynb")
```

# PCE 02 &ndash; Strings


## Q1: Grab a character at a given position in a five-letter string

Write a function called `index_string` that has two parameters:
1. `s` &ndash; a five-letter string;
2. `psn`&ndash; an integer between 0 and 4. 

And that returns:

`result` &ndash; the character that is at position `psn` in `s`.

For example, `index_string("hello", 3)` should return `"l"`.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-fd0dc0c592eb57c1, points: 0, grade: false, solution: true -->
```python
def index_string(s, psn):
    ...

# Test your code
test_s = "shore"
test_psn = 2
print(index_string(test_s, test_psn))
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q1")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q2: Convert a string to all lower- or all upper-case, depending on the input
Write a function called `change_case()` that has two parameters:
1. `string` &ndash; an input string to transform;
2. `direction` &ndash; a string that tells you which direction to change the case (`'lower'` or `'upper'`). 

And that returns:

`new_string` &ndash; a new lower- or upper-case string depending on the direction.

*Hint*: You'll need an if/else conditional for this problem.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-35cf559ffe28dc64, points: 0, grade: false, solution: true -->
```python
def change_case(s, direction):
    ...
    
    
# Test your code
test_s = "sarah"
test_direction = "lower"
change_case(test_s, test_direction)
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q2")
```

</br>
</br>
</br>
</br>
</br>
</br>

---


## Q3: Get domain name from website name


Write a function `get_website_domain()` that has a single parameter:
1. `website` &ndash; a website name (e.g., `'www.umd.edu'`)

And returns:

`domain` &ndash; the domain of the website name, which is the substring _between_ the two `.` characters. 

For example, if I call `get_website_domain("www.umd.edu")`, I should get back `'umd'`, and if I call `get_website_domain("www.cnn.com")`, I should get back `'cnn'`).

__Hint__: This is the split pattern we discussed in class.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-64ee35227074ed60, points: 0, grade: false, solution: true -->
```python
def get_website_domain(website):
    ...


# Test your code
test_website = "www.umd.edu"
print(get_website_domain(test_website))
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q3")
```

</br>
</br>
</br>
</br>
</br>
</br>

---


## Manually graded: Add a docstring to `get_website_domain`

Go back to the solution cell above, and add a docstring to the `get_website_domain` function. The docstring should contain:

1. A brief description of what the function does;
2. A list of the parameters with a brief description of each parameter and its type;
3. A brief description of the return value.

For this task you are free to format the docstring however you like. If you are unsure on how to get started, refer to the &ldquo;Functions&rdquo; slides on ELMS for an example of a docstring that includes all the required pieces of information.

__NOTE__: _There is no autograder for this task: we will grade it manually. Note that you can earn points for this task even if your solution for the previous task is not passing the autograder, however we will check that your docstring does not generate syntax errors of its own._

</br>
</br>
</br>
</br>
</br>
</br>

---

## Q4: Construct a message with string formatting for a waiter splitting the bill for a table

Write a function `split_check()` with parameters:
1. `customer_name` &ndash; the name of the customer
2. `total_check` &ndash; the amount that is owed for the whole table
3. `fraction_owed` &ndash; the fraction of the total check owed by the customer. Should be a flat between 0 and 1 (e.g., `0.2` to denote that this person owes 20% of the check).

And returns the message __"(customer_name) owes (amount_owed)"__, with string formatting so that the `amount_owed` is only showing 2 decimal places. 

* Your program should first compute the amount owed, which should be the product of `total_check` and `fraction_owed`.

* Then it should construct a message that has appropriate string formatting variables for the `customer_name` and `amount_owed` values.

* Finally, it should return this message string.

Note that sometimes the result will be more than 2 decimal places; even in those cases, your resulting message should only display 2 decimal points. For example, `split_check("Joel", 32.50, 0.25)` should return the message "Joel owes 8.12" (even though $32.50 \times 0.25 = 8.125$).

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-710c646983dc4dc3, points: 0, grade: false, solution: true -->
```python
def split_check(customer_name, total_check, fraction_owed):
    ...


# Test your code
test_name = "Joel"
test_check = 32.50
test_fraction = 0.25
print(split_check(test_name, test_check, test_fraction))
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q4")
```

### **_NOTE:_**  Do not forget to run this cell

```python
!jupyter nbconvert --to script PCE-02-Strings.ipynb
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True, run_tests=True, files=['PCE-02-Strings.py'])
```

Tests for q1 (points: None)
```
>>> assert type(index_string("asken", 3)) == str, "Error: return value must be a string"
```
```
>>> assert index_string("asken", 3) == "e", "Wrong return value"
```
```
>>> assert index_string("asken", 1) == "s", "Wrong return value"
```
```
>>> assert index_string("hello", 2) == "l", "Wrong return value"
```
```
>>> assert index_string("hello", 4) == "o", "Wrong return value"
```
```
>>> assert index_string("unify", 0) == "u", "Wrong return value"
```
```
>>> assert index_string("unify", 3) == "f", "Wrong return value"
```


Tests for q2 (points: None)
```
>>> assert type(change_case("sarah", "upper")) == str, "Error: return value must be a string"
```
```
>>> assert change_case("hello", "upper") == "HELLO", "Wrong return value"
```
```
>>> assert change_case("yEs", "lower") == "yes", "Wrong return value"
```
```
>>> assert change_case("BRIDGE", "lower") == "bridge", "Wrong return value"
```
```
>>> assert change_case("hoUSE", "upper") == "HOUSE", "Wrong return value"
```
```
>>> assert change_case("BRIDGE", "upper") == "BRIDGE", "Wrong return value"
```


Tests for q3 (points: None)
```
>>> assert type(get_website_domain("www.umd.edu")) == str, "Error: return value must be a string"
```
```
>>> assert get_website_domain("www.umd.edu") == "umd", "Wrong return value"
```
```
>>> assert get_website_domain("www.cnn.com") == "cnn", "Wrong return value"
```
```
>>> assert get_website_domain("en.wikipedia.org") == "wikipedia", "Wrong return value"
```
```
>>> assert get_website_domain("state.md.gov") == "md", "Wrong return value"
```


Tests for q4 (points: None)
```
>>> assert type(split_check("Joel", 32.50, 0.25)) == str, "Error: return value must be a string"
```
```
>>> assert split_check("joel", 32.5, 0.25) == "joel owes 8.12", "Wrong return value"
```
```
>>> assert split_check("sarah", 21.5, 0.13) == "sarah owes 2.79", "Wrong return value"
```
```
>>> assert split_check("rk", 2.5, 0.79) == "rk owes 1.98", "Wrong return value"
```
```
>>> assert split_check("samar", 7.22, 0.28) == "samar owes 2.02", "Wrong return value"
```
```
>>> assert split_check("joel", 102.8, 0.35) == "joel owes 35.98", "Wrong return value"
```
