# ASSIGNMENT CONFIG
export_cell:
    pdf: false
    run_tests: true
    force_save: true
generate: 
    filtering: true
    zips: false
    show_stdout: true

# PCE 02 &ndash; Strings [Total: 5 pts]

## Q1: Grab a character at a given position in a five-letter string [1 Point]

Write a function called `index_string` that has two parameters:
1. `s` &ndash; a five-letter string;
2. `psn`&ndash; an integer between 0 and 4. 

And that returns:

`result` &ndash; the character that is at position `psn` in `s`.

For example, `index_string("hello", 3)` should return `"l"`.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q1

# BEGIN SOLUTION

<!-- grade_id: cell-fd0dc0c592eb57c1, points: 0, grade: false, solution: true -->
```python
def index_string(s, psn):
    # BEGIN SOLUTION
    return s[psn]
# END SOLUTION
```

# END SOLUTION

Test your function below. Feel free to change the values.

```python
test_s = "shore"
test_psn = 2
print(index_string(test_s, test_psn))
```

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(index_string("asken", 3)) == str, "Error: return value must be a string"
```

<!-- grade_id: cell-ab4e4fc1c1623dc5, points: 1, grade: true, solution: false -->
```python
assert index_string("asken", 3) == "e", "Wrong return value"
```

```python
assert index_string("asken", 1) == "s", "Wrong return value"
```

```python
assert index_string("hello", 2) == "l", "Wrong return value"
```

```python
assert index_string("hello", 4) == "o", "Wrong return value"
```

```python
assert index_string("unify", 0) == "u", "Wrong return value"
```

```python
assert index_string("unify", 3) == "f", "Wrong return value"
```

# END TESTS

# END QUESTION

## Q2: Convert a string to all lower- or all upper-case, depending on the input [1 pt]
Write a function called `change_case()` that has two parameters:
1. `string` &ndash; an input string to transform;
2. `direction` &ndash; a string that tells you which direction to change the case (`'lower'` or `'upper'`). 

And that returns:

`new_string` &ndash; a new lower- or upper-case string depending on the direction.

*Hint*: You'll need an if/else conditional for this problem.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q2

# BEGIN SOLUTION

<!-- grade_id: cell-35cf559ffe28dc64, points: 0, grade: false, solution: true -->
```python
def change_case(s, direction):
    # BEGIN SOLUTION
    if direction == "upper":
        new_string = s.upper()
        return new_string
    elif direction == "lower":
        new_string = s.lower()
        return new_string
# END SOLUTION
```

# END SOLUTION

Test your function below. Feel free to change the values.

```python
test_s = "sarah"
test_direction = "lower"
change_case(test_s, test_direction)
```

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(change_case("sarah", "upper")) == str, "Error: return value must be a string"
```

<!-- grade_id: cell-5bf3c1bb4dc3914a, points: 1, grade: true, solution: false -->
```python
assert change_case("hello", "upper") == "HELLO", "Wrong return value"
```

```python
assert change_case("yEs", "lower") == "yes", "Wrong return value"
```

```python
assert change_case("BRIDGE", "lower") == "bridge", "Wrong return value"
```

```python
assert change_case("hoUSE", "upper") == "HOUSE", "Wrong return value"
```

```python
assert change_case("BRIDGE", "upper") == "BRIDGE", "Wrong return value"
```

# END TESTS

# END QUESTION

## Q3: Get domain name from website name

### Q3.1 Coding Task [1 pt]

Write a function `get_website_domain()` that has a single parameter:
1. `website` &ndash; a website name (e.g., `'www.umd.edu'`)

And returns:

`domain` &ndash; the domain of the website name, which is the substring _between_ the two `.` characters. 

For example, if I call `get_website_domain("www.umd.edu")`, I should get back `'umd'`, and if I call `get_website_domain("www.cnn.com")`, I should get back `'cnn'`).

__Hint__: This is the split pattern we discussed in class.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q3

# BEGIN SOLUTION

<!-- grade_id: cell-64ee35227074ed60, points: 0, grade: false, solution: true -->
```python
def get_website_domain(website):
    # BEGIN SOLUTION
    """Get Website Domain

    Args:
        website (string): Website URL

    Returns:
        string: Website domain
    """
    elements = website.split(".")
    domain = elements[1]
    return domain
# END SOLUTION
```

# END SOLUTION

Test your function below. Feel free to change the values.

```python
test_website = "www.umd.edu"
print(get_website_domain(test_website))
```

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(get_website_domain("www.umd.edu")) == str, "Error: return value must be a string"
```

<!-- grade_id: cell-bc4158632af0431c, points: 1, grade: true, solution: false -->
```python
assert get_website_domain("www.umd.edu") == "umd", "Wrong return value"
```

```python
assert get_website_domain("www.cnn.com") == "cnn", "Wrong return value"
```

```python
assert get_website_domain("en.wikipedia.org") == "wikipedia", "Wrong return value"
```

```python
assert get_website_domain("state.md.gov") == "md", "Wrong return value"
```

# END TESTS

# END QUESTION

## Q3.2: Documentation task [1 pt]

Go back to the solution cell above, and add documentation to the `get_website_domain` function. Your documentation should consist of two parts:
* A docstring [1 pt];
* ⚠️ _NEW_ ⚠️ Line comments [1 pt]. 

The **docstring** should contain:

1. A brief description of what the function does;
2. A list of the parameters with a brief description of each parameter and its type;
3. A brief description of the return value.

For this task you are free to format the docstring however you like. If you are unsure on how to get started, refer to the &ldquo;Functions&rdquo; slides on ELMS for an example of a docstring that includes all the required pieces of information.

The **line comments** should describe concisely what each line of code is doing. Each comment line should start with the `#` character. (For example: `# Initialize return list`.)

__NOTE__: _There is no autograder for this task: we will grade it manually. Note that you can earn points for this task even if your solution for the previous task is not passing the autograder, however we will check that your docstring and comments do not generate syntax errors of their own._

⬇️⬇️⬇️ You can use the cell below to see what your docstring looks like. ⬇️⬇️⬇️

```python
help(get_website_domain)
```

## Q4: Construct a message with string formatting for a waiter splitting the bill for a table [1 Point]

Write a function `split_check()` with parameters:
1. `customer_name` &ndash; the name of the customer
2. `total_check` &ndash; the amount that is owed for the whole table
3. `fraction_owed` &ndash; the fraction of the total check owed by the customer. Should be a flat between 0 and 1 (e.g., `0.2` to denote that this person owes 20% of the check).

And returns the message __"(customer_name) owes (amount_owed)"__, with string formatting so that the `amount_owed` is only showing 2 decimal places. 

* Your program should first compute the amount owed, which should be the product of `total_check` and `fraction_owed`.

* Then it should construct a message that has appropriate string formatting variables for the `customer_name` and `amount_owed` values.

* Finally, it should return this message string.

Note that sometimes the result will be more than 2 decimal places; even in those cases, your resulting message should only display 2 decimal points. For example, `split_check("Joel", 32.50, 0.25)` should return the message "Joel owes 8.12" (even though $32.50 \times 0.25 = 8.125$).

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q4

# BEGIN SOLUTION

<!-- grade_id: cell-710c646983dc4dc3, points: 0, grade: false, solution: true -->
```python
def split_check(customer_name, total_check, fraction_owed):
    # BEGIN SOLUTION
    to_pay = total_check * fraction_owed
    msg = f"{customer_name} owes {to_pay:.2f}"
    return msg
# END SOLUTION
```

# END SOLUTION

Test your function below. Feel free to change the values.

```python
test_name = "Joel"
test_check = 32.50
test_fraction = 0.25
print(split_check(test_name, test_check, test_fraction))
```

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(split_check("Joel", 32.50, 0.25)) == str, "Error: return value must be a string"
```

<!-- grade_id: cell-9106167f5e3e8500, points: 1, grade: true, solution: false -->
```python
assert split_check("joel", 32.5, 0.25) == "joel owes 8.12", "Wrong return value"
```

```python
assert split_check("sarah", 21.5, 0.13) == "sarah owes 2.79", "Wrong return value"
```

```python
assert split_check("rk", 2.5, 0.79) == "rk owes 1.98", "Wrong return value"
```

```python
assert split_check("samar", 7.22, 0.28) == "samar owes 2.02", "Wrong return value"
```

```python
assert split_check("joel", 102.8, 0.35) == "joel owes 35.98", "Wrong return value"
```

# END TESTS

# END QUESTION