# ASSIGNMENT CONFIG
export_cell:
    pdf: true
    run_tests: true
    force_save: true
generate: 
    filtering: true
    zips: false
    show_stdout: true

# Q1: Pad a string with a single-character string

Define a function called `padding` that has two parameters &mdash; a base string, and another string made of a single character.

The function should use a loop to concatenate the character repeatedly to the base string until it reaches a length of 5 (i.e., it _pads_ the base string with the character), and then return the resulting modified string. If the base string is already at length 5 or more it should be returned as it is without any padding.

For example, given `"aa"` and `"b"` as arguments, your function should return `"aabbb"`. If instead you pass in a string that is already 5 characters long (e.g., `abcde`), you should get the same string out of the function.

## Parameters
* `base` &ndash; a string,
* `c` &ndash; single character that needs to be concatenated with the input string.

## Returns
* `newbase` &ndash; new string that has the character `c` added to `base` as many times as it takes to make it 5 characters long.

## Hint
Use an indefinite loop to concatenate the padding character until the desired length has been reached. To check the lenght of a string, you can use the `len()` function.

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q1

# BEGIN SOLUTION

<!-- grade_id: cell-4c9a4b1f07c01a92, points: 1, grade: true, solution: true -->
```python
def padding(base, c):
# BEGIN SOLUTION
    newbase = str(base)
    while len(newbase) < 5:
        newbase += c
    return newbase
# END SOLUTION

# Test with different arguments!
_base = "vvv"
_c = "1"
_newbase = padding(_base, _c)
print("Original:", _base, "length:", len(_base))
print("Padded:", _newbase, "length:", len(_newbase))
```

Output:
```
Original: vvv length: 3
Padded: vvv11 length: 5
```

# END SOLUTION

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(padding("aa", "b")) == str, "For input 'aa' and 'b'', the function should return a str value"
```

```python
assert padding("aa", "b") == "aabbb", "Wrong return value"
```

```python
assert padding("c", "a") == "caaaa", "Wrong return value"
```

```python
assert padding("12345", "b") == "12345", "Wrong return value"
```

# END TESTS

<!-- grade_id: cell-27faca6660a484ee, points: 1, grade: true, solution: false -->
# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# Q2: Multiply numbers in a list

Write a function called `multiplynums` with two parameters &ndash; a list of numbers and a multiplication factor (another number). The function should return a new list where each element has been multiplied by the multiplication factor.

For example, given `[3, 2, 6]` and `2` as arguments, your function should return `[6, 4, 12]`.

## Parameters
- `nums` &ndash; list of numbers
- `multiplier` &ndash; a number

## Returns
- `newnums` &ndash; new list of numbers

---

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q2

# BEGIN SOLUTION

<!-- grade_id: cell-0b0623ccd2463c19, points: 0, grade: false, solution: true -->
```python
def multiplynums(nums, multiplier):
# BEGIN SOLUTION
    newnums = []
    for num in nums:
        product = num * multiplier
        newnums.append(product)
    return newnums
# END SOLUTION

# Test with different arguments!
my_nums = [12, 54, 75, 88, 54]
my_multiplier = 2
my_newnums = multiplynums(my_nums, my_multiplier)
print("Original:", my_nums)
print("Multiplied:", my_newnums)
```

Output:
```
Original: [12, 54, 75, 88, 54]
Multiplied: [24, 108, 150, 176, 108]
```

# END SOLUTION

# BEGIN TESTS

```python
assert type(multiplynums([1, 2, 3, 4, 5], 5)) == list, "For input ([1, 2, 3, 4, 5], 5), the function should return a list"
```

```python
assert multiplynums([1, 2, 3, 4, 5], 5) == [5,10,15,20,25], "Wrong return value"
```

```python
assert multiplynums([0, 1, 2, 3], 4) == [0, 4, 8, 12], "Wrong return value"
```

```python
assert multiplynums([-2, 4, -6], 3)==[-6, 12, -18], "Wrong return value"
```

```python
_original = [1, 2, 3]
_newnums = multiplynums(_original, 3)
assert _original == [1, 2, 3], "Wrong return value"
```

<!-- grade_id: cell-9f0fea7329d1fbd8, points: 1, grade: true, solution: false -->
```python
_original = [4, 5, 6]
_newnums = multiplynums(_original, 5)
assert _original == [4, 5, 6], "Wrong return value"

```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# Q3: Filter a list of numbers that are not even


Define a function called `onlyeven` that takes a list of numbers and returns a list of all even numbers in that original list and no odd number.

For example, with `[2, 4, 65, 67, 89, 87, 88, 90, 23, 34]` as argument, your function should return `[2, 4, 88, 90, 34]`.

### Parameters
* `numbers` &ndash; list of numbers.

### Returns
* `evens` &ndash; a new list with the even numbers that are in `nums`.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q3

# BEGIN SOLUTION

<!-- grade_id: cell-aeefa9518c89deaa, points: 0, grade: false, solution: true -->
```python
def onlyeven(numbers):
# BEGIN SOLUTION
    """
    Filter even numbers from list of numbers
    
    Parameters
    ==========
    numbers : list
        A list of numbers
        
    Returns
    =======
    evens : list
        A new, filtered list with only even numbers
    """
    # Initialize return list
    evens = []
    
    # Loop through each number num in numbers
    for num in numbers:
        
        # Test if num is even
        if num % 2 == 0:
            
            # If that is the case, append it to return list
            evens.append(num)
    
    # Return final list
    return evens
# END SOLUTION

# Test with different arguments!
_numbers = [2, 4, 65, 67, 89, 87, 88, 90, 23, 34]
_evens = onlyeven(_numbers)
print("Original:", _numbers)
print("Filtered:", _evens)
```

Output:
```
Original: [2, 4, 65, 67, 89, 87, 88, 90, 23, 34]
Filtered: [2, 4, 88, 90, 34]
```

# END SOLUTION

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(onlyeven([1, 3, 4, 6, 8, 10])) == list, "For input ([1, 3, 4, 6, 8, 10]), the function should return a list"
```

```python
assert onlyeven([1, 3, 4, 6, 8, 10]) == [4, 6, 8, 10], "Wrong return value"
```

```python
assert onlyeven([1, 3, 5, 7, 9]) == [], "Wrong return value"
```

```python
assert onlyeven([2, 4, 66, 88, 44]) == [2, 4, 66, 88, 44], "Wrong return value"
```

# END TESTS

# END QUESTION

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Manually graded: Add a docstring to `onlyeven`

Go back to the solution cell above, and add a docstring to the `onlyeven` function. The docstring should contain:

1. A brief description of what the function does;
2. A list of the parameters with a brief description of each parameter and its type;
3. A brief description of the return value.

For this task you are free to format the docstring however you like. If you are unsure on how to get started, refer to the &ldquo;Functions&rdquo; slides on ELMS for an example of a docstring that includes all the required pieces of information.

__NOTE__: _There is no autograder for this task: we will grade it manually. Note that you can earn points for this task even if your solution for the previous task is not passing the autograder, however we will check that your docstring does not generate syntax errors of its own._

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

# Q4: Counting odds

Define a function called `howmanyodds` that takes a list of numbers and returns how many odd numbers are in the list. 

For example, given `[2, 4, 65, 67, 89, 87, 88, 90, 22, 34]` as argument, your function should return `3`

## Inputs
* `nums` &ndash; a list of numbers

## Returns
* `count_odds` &ndash; count of odd numbers in the list

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

# BEGIN QUESTION
name: q4

# BEGIN SOLUTION

<!-- grade_id: cell-ee571907f08817b4, points: 0, grade: false, solution: true -->
```python
def howmanyodds(nums):
# BEGIN SOLUTION
    count_odds = 0
    for i in nums:
        if i % 2 != 0:
            count_odds += 1
    return count_odds
# END SOLUTION

# Test with different arguments!
_nums = [2, 4, 65, 67, 89, 87, 88, 90, 23, 34]
_count = howmanyodds(_nums)
print("Original:", _nums)
print("Odds:", _count)
```

Output:
```
Original: [2, 4, 65, 67, 89, 87, 88, 90, 23, 34]
Odds: 5
```

# END SOLUTION

When you're ready run the cell below to get feedback on your answer.

# BEGIN TESTS

```python
assert type(howmanyodds([1, 3, 4, 6, 8, 10])) == int, "For input ([1, 3, 4, 6, 8, 10]), the function should return an int"
```

```python
assert howmanyodds([1, 3, 4, 6, 8, 10]) == 2, "Wrong return value"
```

```python
assert howmanyodds([1, 3, 5, 7, 9]) == 5, "Wrong return value"
```

```python
assert howmanyodds(nums=[2, 4, 66, 88, 44]) == 0, "Wrong return value"
```

# END TESTS

# END QUESTION