```python
# Initialize Otter
import otter
grader = otter.Notebook("PCE-03-Dictionaries.ipynb")
```

# PCE 3: Dictionaries


<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---

## Q1: Create a dictionary

Write a function called `make_dict()` that takes two paramenters &ndash; k, v &ndash; and returns a dictionary with value v corresponding to key k.

Parameters:
- `k` &ndash; the key
- `v` &ndash; the value

Returns:
- `new_dict` &ndash; a dictionary

For example, given `'a'` and `4`as input arguments, `make_dict()` should return `{'a': 4}`

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-7294e987de67cc3d, points: 0, grade: false, solution: true -->
```python
def make_dict(k, v):
...

# Test your function below. Feel free to change the values.
new_dict = make_dict(k='a', v=4)
new_dict
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q1")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q2: Create a dictionary from lists
Write a function called `make_dict_from_lists()` that takes two lists as parameters &ndash; `keys` and `values` &ndash; and returns a dictionary where each element in `keys` is mapped into the corresponding element (by position) in `values`.

Parameters:
- `keys` &ndash; list of elements to use as keys
- `values` &ndash; list of values to use as values

Returns:
- `new_dict` &ndash; a dictionary

For example, given `['a', 'b']`and `['c', 'd']` as arguments, `make_dict_from_lists()` should return `{'a': 'c', 'b': 'd'}`

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-9753350ff6a6c6dc, points: 0, grade: false, solution: true -->
```python

def make_dict_from_lists(keys, values):
...

#Test your function below. Feel free to change the values.

newdict = make_dict_from_lists(keys=['a','c'],values= ['b', 'd'])
newdict
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q2")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q3: Access items from a dictionary

Write a function called `get_details()` that takes a dictionary called `d` and a key called `key` as parameters, and returns the value corresponding to the entered `key` in the dictionary.

Parameters:
- `d` &ndash; a dictionary with 2 or more key&ndash;value pairs
- `key` &ndash; key for which we want extract the corresponding value from the dictionary

Returns:
- `value` &ndash; value corresponding to the key

For example, given as argument this dictionary
```python
dict1 = {
    'laptop_brand': 'Apple',
    'laptop_model': 'MacBook Air',
    'laptop_os' : 'macOS',
    'laptop_os_version': '10.15.7'
}
``` 
 and `"laptop_model"` as the key, `get_details()` should return `"MacBook Air"`
 
 ⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-a4ce783edf9231cd, points: 0, grade: false, solution: true -->
```python
def get_details(d, key):
...

# test your function here
laptop_dict = {
    'laptop_brand': 'Apple',
    'laptop_model': 'MacBook Air',
    'laptop_os' : 'macOS',
    'laptop_os_version': '10.15.7'
}
value = get_details(d = laptop_dict, key = 'laptop_model')
value
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q3")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q4: Add a new entry to a dictionary

Write a function called `add_entry()` that takes one dictionary, a key, and a value and returns the dictionary with the newly added key&ndash;value pair.

Parameters:
- `dict1` &ndash; a dictionary 
- `key` &ndash; a key
- `value` &ndash; a value  

Returns:
- `dict2` &ndash; the original dictionary with the new key&ndash;value pair

For example, given as argument this dictionary
```python
{
    'laptop_brand': 'Apple',
    'laptop_model': 'MacBook Air'
}
``` 
and `'laptop_os'` as the key and `'macOS'` as the value, `add_entry()` should return 
```python
{
    'laptop_brand': 'Apple',
    'laptop_model': 'MacBook Air',
    'laptop_os' : 'macOS'
}
```

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-17ee6d63362e0d09, points: 0, grade: false, solution: true -->
```python

def add_entry(dict1, key, value):
...

# test your function
laptop_dict2 = {
    'laptop_brand': 'Apple',
    'laptop_model': 'MacBook Air'
}

add_entry(dict1= laptop_dict2, key = 'laptop_os', value = 'macOS',)
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q4")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q5: Update value of a key in a dictionary 

Write a function called `update_value_dict()` that takes as parameters a dictionary, an existing key to be updated, and a new value to assign to the key, and returns the updated dictionary.

Parameters:
- `dict1` &ndash; dictionary with 2 or more key&ndash;value pairs
- `key` &ndash; existing key whose value needs to be changed
- `value` &ndash; new value to assign to the existing key

Returns:
- `dict2` &ndash; updated dictionary

For example, given as arguments this dictionary:

```python
dict1 = {
     "INST201": "A-",
     "INST126": "F",
     "INST311": "B+"
}
```

as well as `'INST126'`, and `'A'` as arguments, `update_value_dict()` should return 

```python
dict2 = {
     "INST201": "A-",
     "INST126": "A",
     "INST311": "B+"
}
```
 
⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-ba9d662ddb847de5, points: 0, grade: false, solution: true -->
```python

def update_value_dict(dict1, key, value):
...

# test your function here
classes = {
     "INST201": "A-",
     "INST126": "F",
     "INST311": "B+"
 }

update_value_dict(dict1 = classes,key = 'INST126',value = 'A')
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q5")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q6: Return length of a dictionary

Write a function called `check_length()` that takes a dictionary as input, and returns its length as output.

Parameters:
- `dict1` &ndash; a dictionary with 2 or more key&ndash;value pairs

Returns:
- `length` &ndash; the length of the input dictionary

For example, given as argument
```python
dict1 = {
     "INST201": "A-",
     "INST126": "F",
     "INST311": "B+"
}
``` 

`check_length()` should return `4`

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-ebb41802c3ec0d46, points: 0, grade: false, solution: true -->
```python
def check_length(dict1):
...

# test your function here
classes2 = {
     "INST201": "A-",
     "INST126": "A",
     "INST311": "B+",
     "INST326": "A-"
 }
check_length(dict1 = classes2)
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q6")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q7: Return keys/values based on parameter

Write a function called `return_keys_or_values()` that takes as parameters a dictionary and a string that can be either `"keys"` or `"values"` and returns a list of either keys or values in the dictionary.

Inputs:
- `dict1` &ndash; dictionary with 2 or more key&ndash;value pairs
- `what`  &ndash; specify what to return &ndash; a list of either keys (argument should be `"keys"`) or values (argument should be `"values"`). 

⚠️⚠️ **WARNING**: *if your function uses different strings for this parameter, the autograder will still fail even if your logic is correct!* ⚠️⚠️

Returns:
- `result` &ndash; list of either keys or values in `dict1`.

For example, given as argument
```python
dict1 = {
    'grain': 'Pasta',
    'protein': 'Chicken',
    'vegetable' : 'Broccoli',
    'sauce': 'Butter'
}
```
calling `return_keys_or_values(dict1, "keys")` should return 
```python
['grain', 'protein', 'vegetable', 'sauce']
```
while calling `return_keys_or_values(dict1, "values")` should return 
```python
['Pasta', 'Chicken', 'Broccoli', 'Butter']
```

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-436229d6df5dbbd3, points: 0, grade: false, solution: true -->
```python
def return_keys_or_values(dict1, what):
...

# test your function 
meal_dict = {
    'grain': 'Pasta',
    'protein': 'Chicken',
    'vegetable' : 'Broccoli',
    'sauce': 'Butter'
}
keys = return_keys_or_values(dict1= meal_dict, what = 'keys')
keys
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q7")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q8: Check if a key is present in a dictionary

Write a function called `check_key()` that takes as parameters a dictionary and a key and returns `True` or `False` based on whether the key is present in the dictionary or not.

Parameters:
- `dict1` &ndash; dictionary with some key&ndash;value pairs
- `key` &ndash; a key to check if it is present in the dictionary or not. Could be any object, but for now, assume it is of type `str`.

Returns:
- result &ndash; a boolean based on whether key is present or not

For example, given as argument this dictionary
```python
dict1 = {
    'grain': 'Pasta',
    'protein': 'Chicken',
    'vegetable' : 'Broccoli',
    'sauce': 'Butter'
}
```
the function call `check_key(dict1, "grain")` should return `True`.

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-0ec4d4d49a7c0969, points: 0, grade: false, solution: true -->
```python
def check_key(dict1, key):
...

#test the function
meal_dict = {
   'grain': 'Pasta',
    'protein': 'Chicken',
    'vegetable' : 'Broccoli',
    'sauce': 'Butter'
}

check_key(dict1 = meal_dict,key = 'grain')
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q8")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q9 : Update value of keys if already present in a dictionary 1

Write a function called `update_if_present()` that takes as parameters a base dictionary with some key&ndash;value pairs and an update dictionary with some other key&ndash;value pairs (with some keys in common to the base dict) and returns an updated base dictionary where if key exists in the base dict, its value is updated to the new value in the update dict.

Parameters:
- `base_dict` &ndash; dictionary with some key&ndash;value pairs
- `updates` &ndash; dictionary with keys to be updated (along with their new values)

Returns:
- `updated_base` &ndash; updated base dictionary

For example, given as arguments the dictionaries
```python
base_dict = {
     "INST201": "A-",
     "INST126": "F",
     "INST311": "B+"
}
``` 
and 
```python
updates = {
     "INST126": "C+",
     "INST326": "A-"
}
``` 
 
the function `update_if_present()` should return the dictionary
```python
updated_dict = {
     "INST201": "A-",
     "INST126": "C+",
     "INST311": "B+"
}
```
In other words, it should update `"INST126"` since it was present, but should not add the key `"INST326"` since it is not present in the base.

_**Hint**: you'll need to iterate through the keys and values of the update dictionary._

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-f1816a6a676421ff, points: 0, grade: false, solution: true -->
```python
def update_if_present(base_dict, updates):
...

#test your function
base_dictionary = {
     "INST201": "A-",
     "INST126": "F",
     "INST311": "B+"
 }
dict_updates = {
     "INST126": "C+",
     "INST326": "A-"
 }
updated_dict = update_if_present(base_dict = base_dictionary, updates = dict_updates)
updated_dict
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q9")
```

<br/>
<br/>
<br/>
<br/>
<br/>
<br/>

---


## Q10: Convert a list of records into a dictionary

Write a function called `recordstodict()` that takes as parameters a list of records that are each composed of a key and value separated by a separator (that is also passed in as input), and returns a dictionary representation of that list of strings, mapping each key to its corresponding value. 

Parameters:
- `records` &ndash; list of records, where each record is a string composed of a key and value separated by a separator  character
- `separator` &ndash; a string of a single character that separates the key and value portions of the records

Returns:
- `recsdict` &ndash; a dictionary that maps each key with its corresponding value

For example, given as arguments the list
```python
records = [
     "stephen curry,warriors",
     "lebron james,lakers",
     "kevin durant,nets",
     "giannis antetokounmpo,bucks"
]
``` 
and the string `","`, the function `recordstodict()` should return 
```python
recsdict = {
     "stephen curry": "warriors",
     "lebron james": "lakers",
     "giannis antetokounmpo": "bucks"
}
```

⬇️⬇️⬇️ Enter your solution in the cell below ⬇️⬇️⬇️

<!-- grade_id: cell-e4318517156b4f90, points: 0, grade: false, solution: true -->
```python
def recordstodict(records, separator):
...

# test your function
records = [
     "stephen curry,warriors",
     "lebron james,lakers",
     "kevin durant,nets",
     "giannis antetokounmpo,bucks"
 ]
sep = ","
d = recordstodict(records=records,separator= sep)
d
```

When you're ready run the cell below to get feedback on your answer.

```python
grader.check("q10")
```

### **_NOTE:_**  Do not forget to run this cell

```python
!jupyter nbconvert --to script PCE-03-Dictionaries.ipynb
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True, files=['PCE-03-Dictionaries.py'])
```

Tests for q1 (points: None)
```
>>> assert type(make_dict('a', 'b')) == dict, 'the function should return a dict value'
```
```
>>> assert make_dict('a', 'b') == {'a': 'b'}, 'Wrong return value'
```
```
>>> assert make_dict(0, 'b') == {0: 'b'}, 'Wrong return value'
```
```
>>> assert make_dict(0, 6) == {0: 6}, 'Wrong return value'
```


Tests for q10 (points: None)
```
>>> assert type(recordstodict(['lunch,bread', 'breakfast,cereal', 'dinner,spaghetti'], ',')) == dict, 'the function should return a str value'
```
```
>>> assert recordstodict(['lunch,bread', 'breakfast,cereal', 'dinner,spaghetti'], ',') == {'lunch': 'bread', 'breakfast': 'cereal', 'dinner': 'spaghetti'}, 'Wrong return value'
```
```
>>> assert recordstodict(['bread-$2.50', 'cereal-$1.25', 'spaghetti-$5.00'], '-') == {'bread': '$2.50', 'cereal': '$1.25', 'spaghetti': '$5.00'}, 'Wrong return value'
```
```
>>> assert recordstodict(['INST126|undergraduate', 'INST326|undergraduate', 'INST802|graduate'], '|') == {'INST126': 'undergraduate', 'INST326': 'undergraduate', 'INST802': 'graduate'}, 'Wrong return value'
```


Tests for q2 (points: None)
```
>>> assert type(make_dict_from_lists(['a', 'c'], ['b', 'd'])) == dict, 'the function should return a str value'
```
```
>>> assert make_dict_from_lists(['a', 'c'], ['b', 'd']) == {'a': 'b', 'c': 'd'}, 'Wrong return value'
```
```
>>> assert make_dict_from_lists([0, 'b'], ['c', 1]) == {0: 'c', 'b': 1}, 'Wrong return value'
```
```
>>> assert make_dict_from_lists([0, 6], [5, 1]) == {0: 5, 6: 1}, 'Wrong return value'
```


Tests for q3 (points: None)
```
>>> assert type(get_details({'smartphone_brand': 'Samsung', 'smartphone_model': 'S20'}, 'smartphone_brand')) == str, 'the function should return a str value'
```
```
>>> assert get_details({'smartphone_brand': 'Samsung', 'smartphone_model': 'S20'}, 'smartphone_brand') == 'Samsung', 'Wrong return value'
```
```
>>> assert get_details({'a': 'b', 'c': 'd'}, 'a') == 'b', 'Wrong return value'
```
```
>>> assert get_details({'p': 'q', 'r': 's'}, 'r') == 's', 'Wrong return value'
```


Tests for q4 (points: None)
```
>>> assert type(add_entry({'l': 'm', 'n': 'o'}, 'p', 'q')) == dict, 'the function should return a dict value'
```
```
>>> assert add_entry({'l': 'm', 'n': 'o'}, 'p', 'q') == {'l': 'm', 'n': 'o', 'p': 'q'}, 'Wrong return value'
```
```
>>> assert add_entry({'a': 'b', 'c': 'd'}, 'p', 'q') == {'a': 'b', 'c': 'd', 'p': 'q'}, 'Wrong return value'
```
```
>>> assert add_entry({9: 'b', 'c': 'd'}, 66, 'q') == {9: 'b', 'c': 'd', 66: 'q'}, 'Wrong return value'
```


Tests for q5 (points: None)
```
>>> assert type(update_value_dict({'a': 'b', 'c': 'd'}, 'a', 'hi')) == dict, 'the function should return a dict value'
```
```
>>> assert update_value_dict({'a': 'b', 'c': 'd'}, 'a', 'hi') == {'a': 'hi', 'c': 'd'}, 'Wrong return value'
```
```
>>> assert update_value_dict({'a': 'b', 'c': 'd'}, 'c', 'hello') == {'a': 'b', 'c': 'hello'}, 'Wrong return value'
```
```
>>> assert update_value_dict({'a': 'b', 'c': 'd'}, 'a', 2) == {'a': 2, 'c': 'd'}, 'Wrong return value'
```


Tests for q6 (points: None)
```
>>> assert type(check_length({'a': 'b', 'c': 'd'})) == int, 'the function should return a int value'
```
```
>>> assert check_length({'a': 'b', 'c': 'd'}) == 2, 'Wrong return value'
```
```
>>> assert check_length({'a': 'b', 'c': 'd', 'e': 'f'}) == 3, 'Wrong return value'
```
```
>>> assert check_length({'a': 'b'}) == 1, 'Wrong return value'
```


Tests for q7 (points: None)
```
>>> assert type(set(return_keys_or_values({'brand': 'Samsung', 'model': 'S20', 'os': 'Android', 'os_version': '10.2.2'}, 'keys'))) == set, 'the function should return a set value'
```
```
>>> assert set(return_keys_or_values({'brand': 'Samsung', 'model': 'S20', 'os': 'Android', 'os_version': '10.2.2'}, 'keys')) == {'brand', 'model', 'os', 'os_version'}, 'Wrong return value'
```
```
>>> assert set(return_keys_or_values({'brand': 'Samsung', 'model': 'S20', 'os': 'Android', 'os_version': '10.2.2'}, 'values')) == {'Samsung', 'S20', 'Android', '10.2.2'}, 'Wrong return value'
```
```
>>> assert set(return_keys_or_values({'a': 'b'}, 'keys')) == {'a'}, 'Wrong return value'
```


Tests for q8 (points: None)
```
>>> assert type(check_key({'smartphone_brand': 'Apple'}, 'smartphone_brand')) == bool, 'the function should return a bool value'
```
```
>>> assert check_key({'smartphone_brand': 'Apple'}, 'smartphone_brand') == True, 'Wrong return value'
```
```
>>> assert check_key({'smartphone_brand': 'Apple'}, 'smartphone_os') == False, 'Wrong return value'
```
```
>>> assert check_key({'a': 'b', 0: 'd'}, 0) == True, 'Wrong return value'
```


Tests for q9 (points: None)
```
>>> assert type(update_if_present({'name': 'Joel', 'age': 22, 'job': 'student'}, {'age': 35, 'eyes': 'brown'})) == dict, 'the function should return a dict value'
```
```
>>> assert update_if_present({'name': 'Joel', 'age': 22, 'job': 'student'}, {'age': 35, 'eyes': 'brown'}) == {'name': 'Joel', 'age': 35, 'job': 'student'}, 'Wrong return value'
```
```
>>> assert update_if_present({'item': 'iPhone', 'brand': 'Apple', 'cost': 300.0}, {'item': 'iPhone Pro', 'color': 'gold'}) == {'item': 'iPhone Pro', 'brand': 'Apple', 'cost': 300.0}, 'Wrong return value'
>>> print('PASS')
PASS
```
```
>>> assert update_if_present({'pocket': 'empty', 'backpack': 'full', 'briefcase': 'empty'}, {'pocket': 'full', 'umbrella': 'open'}) == {'pocket': 'full', 'backpack': 'full', 'briefcase': 'empty'}, 'Wrong return value'
>>> print('PASS')
PASS
```
